---
title: 'EEG Analysis'
description: ""
---

## Introduction to EEGs

```yaml
type: NormalExercise
key: f9ed121817
xp: 100
```

[Electroencephalography](https://en.wikipedia.org/wiki/Electroencephalography) is an electrophysiological monitoring method to record electrical activity of the brain. The human body transfers information by neurons. This cells emitting singals via changes of electrochemical potentials. Invivo normally a single neuronal activity can't detected. But in thought process mostly hundreds of cells are involved, why we can measure electrical potentials on the scalp. In this exercise we will focus on 2 EEG channels.

The signals measured in milli volts (mv) and need some simple pre-processing to obtain information. At the end of the exercise we can compare $\alpha$-, $\beta$- and $\delta$-bands. 

$\alpha$-Wellen (8-13 Hz):
- entspannter Wachzustand, i.b. bei geschlossenen Augen

$\beta$-Wellen (>13 Hz):
- während mentaler oder körperlicher Aktivität oder bei psychischer Belastung

$\delta$-Wellen (0,5-3 Hz):
- im Tiefschlaf

`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```
