---
  title: "Task 1"
  description: "First Task"
---

## An exercise title written in sentence case

```yaml
type: NormalExercise 
lang: r
xp: 100 
skills: 1
key: 847996f3f9   
```


This is a first task to get familar with R and time series analysis.
We wanna start with some easy tasks. For time series analysis we need some data. So let's create some random numbers. Therefor we can use the command `[runif()](https://www.rdocumentation.org/packages/compositions/versions/1.40-2/topics/runif)`


`@instructions`
- Instruction 1
- Instruction 2
- Instruction 3
- Instruction 4

`@hint`
- Here is the hint for this setup problem. 
- It should get students 50% of the way to the correct answer.
- So don't provide the answer, but don't just reiterate the instructions.
- Typically one hint per instruction is a sensible amount.

`@pre_exercise_code`

```{r}
# Load datasets and packages here.
```

`@sample_code`

```{r}
# Your
# sample
# code
# should
# be
# ideally
# 10 lines or less,
# with a max
# of 16 lines.
```

`@solution`

```{r}
# Answer goes here
# Make sure to match the comments with your sample code
# to help students see the differences from solution
# to given.
```

`@sct`

```{r}
# Update this to something more informative.
success_msg("Some praise! Then reinforce a learning objective from the exercise.")
```
